# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your GT login and the GT logins of any of your collaborators below. (The GT logins are worth 1 point per notebook, so don't miss the opportunity to get a free point!)

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Lab 0-B: Linear models [6 points]

We say $F(x_{t-1})$ is _linear_ if it may be written in the form,

$$
  F(x_{t-1}) \equiv m x_{t-1} + b,
$$

where $m$ and $b$ are constants.

Let's see how this model behaves as $m$ and $b$ vary.

```python
from random import randint

%matplotlib inline
import matplotlib.pyplot as plt
```

**Exercise 1** (1 point). Define a linear autonomous first-order state transition function, `F_linear(x, m=1, b=0)`, where the `m` and `b` arguments correspond to the constant coefficients, $m$ and $b$, which take on default values of $1$ and $0$, respectively.

<!-- grade_id: F_linear, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: F_linear_test, points: 1, grade: true, solution: false -->
```python
from random import randint

for trial in range (10):
    m = randint (-10, 10)
    b = randint (-10, 10)
    x = randint (-10, 10)
    y = F_linear (x, m, b)
    assert (y - b) == (x*m)
    
print ("\n(Passed!)")
```

**Exercise 2** (2 points). Complete the simulator below. It is a function, `sim`, that takes the following arguments.

- `fun`: A function that implements the autonomous first-order state transition function of a discrete dynamical system. It should have the calling signature, `fun(x, ...)`, where `x` is the current state and `...` can accept any additional (and optional) named arguments. See `fun_args` below.
- `t_max`: The number of time steps to simulate.
- `x0`: The value of the initial state.
- `fun_args`: Any optional additional named arguments to be passed to `fun`.

It should return a list `X[:t_max+1]` where `X[t]` is the state variable at time step `t`. Element `X[0] == x0`.

<!-- grade_id: sim, points: 0, grade: false, solution: true -->
```python
def sim (fun, t_max, x0, **fun_args):
    X = [x0]
    for t in range (1, t_max+1):
        # YOUR CODE HERE
        raise NotImplementedError()
    return X
```

<!-- grade_id: sim_test, points: 2, grade: true, solution: false -->
```python
T_MAX = 30
M = 1.1
B = 0
X0 = 1.0

X = sim (F_linear, T_MAX, X0, m=M, b=B)
T = list (range (T_MAX+1))
plt.plot (T, X, 's')

def f_lin_det (t, x0, m, b):
    c = b / (m-1)
    return (m**t)*(x0 + c) - c

Y = [f_lin_det (t, X0, M, B) for t in T]
plt.plot (T, Y, 'r-')

import numpy as np
MACH_EPS = np.finfo (float).eps

rdy = [(y-x)/x for x, y in zip (X, Y)]
assert all ([(abs (r) <= 2.*MACH_EPS) for r in rdy])

print ("\n(Passed!)")
```

**Exercise 3** (3 points). For autonomous, first-order, discrete linear dynamical system, $x_t = m x_{t-1} + b$, show that

$$
  x_t = m^t (x_0 + c) - c,
$$

where $c \equiv \dfrac{b}{m-1}$.

> You'll notice that the preceding testing code cell above compares the output of your simulator with a collection of values computed using this deterministic formula.

<!-- grade_id: proof, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```