# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your GT login and the GT logins of any of your collaborators below. (The GT logins are worth 1 point per notebook, so don't miss the opportunity to get a free point!)

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Lab 0-A: Autonomous First-Order Discrete Dynamical Model

Our first major modeling concept is that of an **autonomous first-order system** that is **discrete in both time and space.**

Formally, let $t = 0, 1, 2, \ldots$ denote time, and let $x_t$ be the system's state variable, which also takes on values from some discrete set $S$. Then the system evolves in time as

$$
  x_t = F(x_{t-1})
$$

where $F(\cdot)$ is the system's _state transition function_. This system is _first-order_ because the state at time $t$ depends only on the state at time $t-1$; it is _autonomous_ because it does not depend explicitly on the variable $t$. (That is, it depends only implicitly on $t$ through $x_{t-1}$.)

## Example: Trend setters and followers

Consider a world in which sweaters come in two colors, black or white. Within this world, _trend setters_ decide what to wear and _trend followers_ mimic the style of trend setters.

Let's model this world using two agents, $A$ and $B$. Agent $A$ is the trend setter and $B$ is the trend follower. Their sweaters come in two colors, 0 (black) or 1 (white). Initially, $A$ wears a black sweater and $B$ wears a white sweater. The two agents change their sweaters according to the following rules:

- If agent $A$ observes that $B$ is wearing a sweater of the same color, then $A$ changes its sweater color.
- Agent $B$ always changes his or her sweater color to match that of $A$.

**Exercise 1** (2 points). Assume $x_t$ is a tuple, $x_t \equiv (a_t, b_t)$, where $a_t$ and $b_t$ denote the colors of $A$ and $B$ at time $t$. Write down the mathematical form of $F(x_{t-1})$. Use a Markdown cell (below) and LaTeX syntax for the mathematical bits.

> Detexify is a neat site for "looking up" a math symbol in LaTeX: http://detexify.kirelabs.org/classify.html

<!-- grade_id: F_math, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 2** (2 points). Implement a state transition function, `F(a, b)`, that takes as input the colors `a` and `b` of agents $A$ and $B$, respectively, at the current time step and returns their updated colors according to the governing rule above. The function should return the new colors as a pair. The valid colors are 0 and 1.

<!-- grade_id: F_code, points: 0, grade: false, solution: true -->
```python
def F (a, b):
    assert a in [0, 1]
    assert b in [0, 1]
    
    a_new, b_new = a, b
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return a_new, b_new
```

<!-- grade_id: F_code_test, points: 2, grade: true, solution: false -->
```python
# Test code
assert F (0, 0) == (1, 0)
assert F (0, 1) == (0, 0)
assert F (1, 0) == (1, 1)
assert F (1, 1) == (0, 1)

print ("\n(Passed!)")
```

**Exercise 3** (2 points). Using your function, `F(a, b)`, complete the following simulation of this system. 

<!-- grade_id: sim, points: 0, grade: false, solution: true -->
```python
t = 0 # Discrete time counter
T_MAX = 10 # Number of simulation steps

# Tracks all state changes over time:
A = [0]
B = [1]

# YOUR CODE HERE
raise NotImplementedError()
    
print ("A:", A)
print ("B:", B)
```

<!-- grade_id: sim_test, points: 2, grade: true, solution: false -->
```python
assert len (A) == len (B)

for i in range (1, len (A)):
    a0, b0 = A[i-1], B[i-1]
    a, b = A[i], B[i]
    assert ((a0 != b0) and (a == b)) or ((a0 == b0) and (a != b))
    
print ("\n(Passed!)")
```

**Exercise 4** (1 point). Plot the result as a time series using Matplotlib's [`plot()` function](http://matplotlib.org/users/pyplot_tutorial.html) (in the `matplotlib.pyplot` module).

<!-- grade_id: plot, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Hey, what do you know? Fashion trends are cyclical.

**Exercise 5** (2 points). Make a phase space plot of this system.

<!-- grade_id: phase_plot, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```