# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your GT login and the GT logins of any of your collaborators below. (The GT logins are worth 1 point per notebook, so don't miss the opportunity to get a free point!)

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Lab 0-C: Oscillations

The linear model only captures exponential growth (toward infinity) or decay (toward some constant). It does not appear to be able to capture other kinds of behavior, such as the simple cyclic behavior we saw in the "trend setting" example.

Indeed, as the trend setting example suggests, we will need multiple variables to express oscillation within the linear framework. Let's see how.

```python
import numpy as np
MACH_EPS = np.finfo(float).eps

import matplotlib.pyplot as plt

%matplotlib inline
```

**Exercise 1** (2 points). Implement the state transition function for the following system of two variables.

$$
\begin{eqnarray}
  x_t & = & \frac{1}{2} x_{t-1} + y_{t-1} \\
  y_t & = & -\frac{1}{2} x_{t-1} + y_{t-1}
\end{eqnarray}
$$

In particular, given an input tuple `z = (x, y)` at some time step $t$, where $x_{t-1}=$`x` and $y_{t-1}=$`y`, it should return a new tuple `(u, v)` where `u`$=x_t$ and `v`$=y_t$.

<!-- grade_id: F, points: 0, grade: false, solution: true -->
```python
def F (z):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: F_test, points: 2, grade: true, solution: false -->
```python
assert F (F (F (F (F ((2, -1)))))) == (-0.75, 1.75)
print ("\n(Passed.)")
```

**Exercise 2** (2 points). Implement a simulation loop for this system. Your implementation should be a function `sim (z0, t_max)` that takes the initial state `z0` and number of time steps `t_max` as inputs, and returns a list `Z[:t_max+1]` such that `Z[t]` is the state of the system at time step `t`, with `Z[0] == z0`.

<!-- grade_id: sim, points: 0, grade: false, solution: true -->
```python
# def sim (z0, t_max):
#     ...

# YOUR CODE HERE
raise NotImplementedError()

T_MAX = 50
Z0 = (1, 1)
Z = sim (Z0, T_MAX)

T = list (range (T_MAX+1))
X = [x for x, _ in Z]
Y = [y for _, y in Z]
plt.plot (T, X, 's-', T, Y, '*--')
```

<!-- grade_id: sim_test, points: 2, grade: true, solution: false -->
```python
Z_test = sim ((1, 1), 50)

X_soln = []
Y_soln = []
with open ('Z_soln__1_1_50.csv', 'rt') as fp:
    for line in fp.readlines ():
        xs, ys = tuple (line.strip ().split (','))
        X_soln.append (float (xs))
        Y_soln.append (float (ys))

X_test = [x for x, _ in Z_test]
dX_test = [x-xs for x, xs in zip (X, X_soln)]
Y_test = [y for _, y in Z_test]
dY_test = [y-ys for y, ys in zip (Y_test, Y_soln)]

assert all ([abs (dx) <= 4.*MACH_EPS for dx in dX_test])
assert all ([abs (dy) <= 4.*MACH_EPS for dy in dY_test])

print ("\n(Passed.)")
```

**Exercise 3** (2 points). This system may be written as a system of linear equations of the form,

$$
  z_t = M z_{t-1},
$$

where $M$ is a matrix. What is $M$ for this example?

<!-- grade_id: M, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```