# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your GT login and the GT logins of any of your collaborators below. (The GT logins are worth 1 point per notebook, so don't miss the opportunity to get a free point!)

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Logistic growth

> These notes follow Sayama's procedure from Chapter 4 of his book, albeit with a different notation.

Recall the simplest first-order autonomous linear system,

$$
  x_t = (1 + \alpha) x_{t-1},
$$

where $x_t$ is the size of some population at time $t$ and $1 + \alpha > 0$ is a parameter that describes the rate of growth or decline of this population. I have written this rate as "$1 + \alpha$" so you can interpret $\alpha$ as a "percent growth" or "percent decay." 

> Typical populations grow or decay at a few percent annually. For instance, I might expect to see $\alpha$ in the range of $\pm \, 3$%, or $\alpha \in [-0.03, +0.03]$.

The solution to this recurrence in terms of $t$ is simply

$$
  x_t = (1 + \alpha)^t x_0.
$$

Regarding the behavior of this model, there are three possibilities.

- $-1 < \alpha < 0$: The population contracts or decays in time toward 0.
- $\alpha > 0$: The population grows exponentially over time without bound.
- $\alpha = 0$: The population remains the same.

This behavior is interesting but somewhat unrealistic. For example, we might expect a real population to grow but only to some limit, $x_t \leq \kappa$, where $\kappa$ is some maximum capacity that the environment can support. How should we modify the model to enforce such a limit?

Following Sayama's tips, let's see if we can hack this model to account for capacity. There are many potential strategies, but a good one is to take an existing parameter, factor, or term and try replacing it with some new function, $a(x)$, i.e.,

$$
  x_t = \left[ 1 + a(x_{t-1}) \right] \cdot x_{t-1}.
$$

In this case, $1 + a(x)$ governs the rate of growth, which should slow down and eventually stop as the population approaches its maximum capacity. That is, as $x_t \rightarrow \kappa$, we want $a(x_t) \rightarrow 0$.

Again, Sayama's advice is to start with a simple form for the unknown function $a(x)$, such as a line.

Recall that two points define a line. In our case, when $x=0$, then $a(x) = \alpha$, which is the "natural" or "intrinsic" growth rate of the original model. And when $x=\kappa$, $a(x) = 0$. Thus,

$$
\begin{eqnarray}
  a(x) & = & (\text{slope}) x + (\text{intercept}) \\
       & = & - \frac{\alpha}{\kappa} x + \alpha \\
       & = & \alpha \left( 1 - \frac{x}{\kappa} \right).
\end{eqnarray}
$$

The new hypothesized growth model is

$$
\begin{eqnarray}
  x_t & \equiv & \left[ 1 + \alpha \left( 1 - \frac{x_{t-1}}{\kappa} \right) \right] \cdot x_{t-1}.
%      &    =   & x_{t-1} + \alpha x_{t-1} - \alpha \frac{x_{t-1}}{\kappa} x_{t-1}.
\end{eqnarray}
$$

Let's try it out!

**Exercise 1** (2 points). Implement a state transition function, $F(x)$, corresponding to the preceding model. That is, given some scalar input value, $x_{t-1}$, your function should return a value $F(x_{t-1}) = x_t$.

<!-- grade_id: F, points: 0, grade: false, solution: true -->
```python
def F (x, alpha, kappa):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: F__test, points: 2, grade: true, solution: false -->
```python
import numpy as np
MACH_EPS = np.finfo (float).eps

assert np.abs (F (20, .025, 50) - 20.3) / 20.3 <= 5*MACH_EPS
assert np.abs (F (100, .05, 50) - 95) / 95 <= 5*MACH_EPS
```

**Exercise 2** (2 points). Implement a simulator for this system. In particular, your simulation should produce a NumPy array, `X[:t_max+1]`, containing all $x_t$ for $t \in \{1, 2, \ldots, $ `t_max` $\}$, given $x_0=$`X[0]`.

<!-- grade_id: sim_F, points: 2, grade: true, solution: true -->
```python
import numpy as np

def sim (F, t_max, x0, **args):
    """
    `sim (F, t_max, x0, **args)`: Given the state transition
    function for a first-order scalar autonomous system, this
    function simulates the system for `t_max` time steps.
    
    It returns a NumPy array, `X[:t_max+1]`, whose first
    entry `X[0]` is the given initial state, `x0`, and whose
    subsequent entries, `X[:, t]` for `1 <= t <= t_max` are
    the system's state at each time step `t`.
    
    The state transition function may accept additional named
    arguments, which can be passed to it via the variable
    keyword parameter dictionary, `**args`.
    """
    X = np.zeros (t_max+1)
    # YOUR CODE HERE
    raise NotImplementedError()
    return X

T_MAX = 200
INIT_POP = 1
ALPHA = .05 # 5% growth per unit time
KAPPA = 100

X = sim (F, t_max=T_MAX, x0=INIT_POP, alpha=ALPHA, kappa=KAPPA)

# Prints first 10 steps and last 10 steps
print (X[:10])
print ("...")
print (X[-10:])
```

```python
from matplotlib.pyplot import plot, suptitle, title, xlabel, axis
%matplotlib inline

plot (X, '.')
plot ([0, T_MAX], [KAPPA, KAPPA], 'r-.')
axis ([-1, T_MAX+1, 0, KAPPA*1.1])

title (r'Population under the logistic growth model: $\alpha={}$, $\kappa={}$'.format (ALPHA, KAPPA))
xlabel (r'Time')
```

This model is known in the research literature as the _logistic growth model_.

**Exercise 3** (1 point). Repeat the above exercise keeping all parameters the same _except_ the initial population, which should start at 250 individuals. (Note that this value exceeds the capacity, `KAPPA=100`.) At what time step $t$ does $x_t$ first fall within 1% of $\kappa=100$? Write your answer in the Markdown cell below (i.e., the one that initially reads, `YOUR ANSWER HERE`).

<!-- grade_id: sim2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: sim2_t_one_percent, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```