# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your GT login and the GT logins of any of your collaborators below. (The GT logins are worth 1 point per notebook, so don't miss the opportunity to get a free point!)

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# A predator-prey system

Consider a population of hippos and lions. The lions are _predators_ and the hippos are _prey_, since the lions eat the hippos but not vice-versa. Suppose we wish to model these populations.

## Part 1: Population dynamics in isolation

To develop a model, let's think first about the _intrinsic_ dynamics of each population separately from the other.

- In the absence of lions, the hippos multiply.
- In the absence of hippos, the lions die.

Mathematically, if $H_t$ is the number of hippos at time $t$ and $L_t$ is the number of wolves at time $t$, then our simple growth and decay models would give us the following starting point.

$$
\begin{eqnarray}
  H_t & \sim & \left[ 1 + \gamma \left( 1 - \frac{H_{t-1}}{\kappa} \right) \right] H_{t-1} \\
  L_t & \sim & (1 - \delta) L_{t-1},
\end{eqnarray}
$$

where $0 \leq \gamma, \delta$ are the percent increase and decrease per unit time of the two populations, respectively; and $\kappa$ is the maximum population capacity of the sheep.

## Part 2: Interactions

Now let's add interaction terms between the two groups.

Start by thinking about the number of hippos that will be eaten in some time period. As before, we might start by supposing that the number of hippos that die due to lions behaves like

$$
  d(L_{t-1}) \cdot H_{t-1},
$$

where $d(l)$ is a function that lies between 0 and 1, inclusive. The case of $d(l)=0$ corresponds to no deaths, which should happen when $l=0$; and $d(l) \rightarrow 1$ when there are many lions, i.e., as $w \rightarrow \infty$.

A simple function with this behavior is,

$$
  d(l) \equiv 1 - \frac{1}{\beta l + 1},
$$

where $\beta > 0$ is some parameter that can be used to control how quickly lions consume hippos.

**Exercise 1** (2 points). Plot $d(l)$ for the following four values of $\beta$: 0.01, 0.1, 1, and 10.

<!-- grade_id: plot_d_of_l, points: 2, grade: true, solution: true -->
```python
from matplotlib.pyplot import plot, legend, xlabel, title
%matplotlib inline

# YOUR CODE HERE
raise NotImplementedError()
```

Similarly, we might expect the lion population to grow faster the more hippos there are:

$$
  L_t \sim \left[ 1 + g(H_{t-1}) \right] L_{t-1}.
$$

The simplest function describing growth might be our usual $g(h) = \eta h$, which measures percent growth (as a _function_ of $h$).

Let's put these ideas together into the following mathematical model:

$$
\begin{eqnarray}
  H_t & \equiv & \left[ 1 + \gamma \left( 1 - \frac{H_{t-1}}{\kappa} \right) - \left( 1 - \frac{1}{\beta L_{t-1} + 1} \right) \right] H_{t-1} \\
  L_t & \equiv & (1 - \delta + \eta H_{t-1}) L_{t-1}.
\end{eqnarray}
$$

Since $d(l) h$ counts the number of the original $h$ hippos who are eaten by the $l$ lions, we subtract it from $H_{t-1}$ in the first equation; and since $\eta h$ measures the percent increase in lions due to the presence of $h$ hippos, it offsets the percent decrease due to the lions natural death rate ($\delta$).

Let's simulate this model for the following parameter settings:
- $\gamma=1$: In isolation the prey naturally double in each time step.
- $\kappa=5$: The environment supports a maximum of five prey.
- $\beta=1$: This parameter does not have an easy interpretation, other than qualitative effect of increasing the death rate for larger $\beta$. In practice, you would calibrate it against observations.
- $\delta=1$: The predators would naturally immediately die out if they were not able to feed within a given time step.
- $\eta=1$: Like $\beta$, this parameter only has the qualitative interpretation previously discussed.

Let's also use the initial conditions of $H_0=L_0=1$ (one hippo and one lion), and simulate for `t_max=100` time steps.

**Exercise 2** (4 points). Simulate this system for the parameters given above. In particular, it should produce its results in the form of a NumPy array, `X[:2, :t_max+1]`, where each column `X[:, t]` is the state of the simulation at time `t` as a vector with two components: `X[0, t]` being the population of hippos at time t and `X[1, t]` being the population of lions.

The code that follows your simulation assumes `X` as input and creates both a time-series plot and a phase-space plot.

<!-- grade_id: HL_sim, points: 4, grade: true, solution: true -->
```python
import numpy as np

# Create a simulation and store your results
# in a NumPy array, `X`, as instructed above.

# YOUR CODE HERE
raise NotImplementedError()

# ==== The following plots your computed `X` ====

# Helper function to create a plot
def plot_sim (X):
    assert type (X) is np.ndarray
    assert len (X.shape) == 2 and X.shape[0] == 2
    
    H = X[0, :]
    L = X[1, :]

    from matplotlib.pyplot import subplots
    f, ax = subplots (1, 2, figsize=(10, 5))
    ax[0].plot (H, 'b-', label='hippos')
    ax[0].plot (L, 'g-.', label='lions')
    ax[0].set_title ('Populations')
    ax[0].legend (loc='best')
    ax[1].plot ([1], [1], 'rs')
    ax[1].plot (H, L, '.-')
    ax[1].set_title ('Phase space')
    ax[1].set_xlabel ('hippos')
    ax[1].set_ylabel ('lions', rotation=0)
    
plot_sim (X)
```

## Capping the lion population

**Exercise 3** (2 points). There are parameters for which the lion population "blows up." Fill in the code cell below to re-run your simulation but for the parameters $\gamma=1$, $\kappa=50$, $\beta=1$, $\delta=1$, and $\eta=1$. (Use the same initial conditions and number of time steps.) Store the result in a NumPy array, `Y`.

<!-- grade_id: blow_up, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

plot_sim (Y)
```

A simple fix might be to cap the overall lion population in the same way we cap the hippo population, namely, by introducing population capacity parameters: $\kappa_h$ for hippos and $\kappa_l$ for lions.

$$
\begin{eqnarray}
  H_t & \equiv & \left[ 1 + \gamma \left( 1 - \frac{H_{t-1}}{\kappa_h} \right) - \left( 1 - \frac{1}{\beta L_{t-1} + 1} \right) \right] H_{t-1} \\
  L_t & \equiv & \left[ 1 + \eta H_{t-1} \left( 1 - \frac{L_{t-1}}{\kappa_l} \right) - \delta \right] L_{t-1}.
\end{eqnarray}
$$

**Exercise 4** (4 points). Simulate this system. Use the same simulation parameters as in Exercise 3, with tne new parameters $\kappa_h=50$ and $\kappa_l=250$. Store the simulation results in a NumPy array, `Z`.

<!-- grade_id: Z, points: 4, grade: true, solution: true -->
```python
# Simulate and produce an output array, `Z`

# YOUR CODE HERE
raise NotImplementedError()

# Plots your `Z`:
plot_sim (Z)
```

**Exercise 5** (3 points). In Exercises 2 and 4, you should have observed that the populations appear to tend toward an equilibrium point. Read about the root-finder, [`scipy.optimize.root()`](https://docs.scipy.org/doc/scipy-0.18.1/reference/generated/scipy.optimize.root.html). Write some code to determine this equilibrium point for the model of Exercise 2 and compute equilibrium point. Store the point as a NumPy array, `x_eq[:2]` where `x_eq[0]` is the hippo equilibrium value and `x_eq[1]` is the lion equilibrium value. Also find the roots of the model of Exercise 4 and store its results in `z_eq[:2]`.

<!-- grade_id: roots, points: 3, grade: true, solution: true -->
```python
from scipy.optimize import root

# Compute `x_eq[:2]` and `z_eq[:2]`
# YOUR CODE HERE
raise NotImplementedError()

# Prints your computed results. Are they consistent with
# the trends of your plots from Exercises 2 and 4?
print ("\nYour x_eq:", x_eq)
print ("Your z_eq:", z_eq)
```

```python

```